{
 "cells": [
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [],
   "source": [
    "import requests"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [],
   "source": [
    "url = \"http://tptp.org/cgi-bin/SystemOnTPTPFormReply\""
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 16,
   "metadata": {},
   "outputs": [],
   "source": [
    "problem = \"\"\"%------------------------------------------------------------------------------\n",
    "% File       : E---2.4\n",
    "% Problem    : PUZ001+1 : TPTP v7.3.0. Released v2.0.0.\n",
    "% Transform  : none\n",
    "% Format     : tptp:raw\n",
    "% Command    : eprover --delete-bad-limit=2000000000 --definitional-cnf=24 -s --print-statistics -R --print-version --proof-object --auto-schedule --cpu-limit=%d %s\n",
    "\n",
    "% Computer   : n120.star.cs.uiowa.edu\n",
    "% Model      : x86_64 x86_64\n",
    "% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz\n",
    "% Memory     : 32218.5MB\n",
    "% OS         : Linux 3.10.0-862.11.6.el7.x86_64\n",
    "% CPULimit   : 300s\n",
    "% DateTime   : Thu Jul 11 20:00:13 EDT 2019\n",
    "\n",
    "% Result     : Theorem 0.03s\n",
    "% Output     : CNFRefutation 0.03s\n",
    "% Verified   : \n",
    "% Statistics : Number of formulae       :   47 (  68 expanded)\n",
    "%              Number of clauses        :   25 (  32 expanded)\n",
    "%              Number of leaves         :   11 (  18 expanded)\n",
    "%              Depth                    :    8\n",
    "%              Number of atoms          :   85 ( 128 expanded)\n",
    "%              Number of equality atoms :   26 (  41 expanded)\n",
    "%              Maximal formula depth    :    5 (   3 average)\n",
    "%              Maximal clause size      :    4 (   2 average)\n",
    "%              Maximal term depth       :    2 (   1 average)\n",
    "\n",
    "% Comments   : \n",
    "%------------------------------------------------------------------------------\n",
    "fof(pel55,conjecture,(\n",
    "    killed(agatha,agatha) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55)).\n",
    "\n",
    "fof(pel55_3,axiom,(\n",
    "    ! [X1] :\n",
    "      ( lives(X1)\n",
    "     => ( X1 = agatha\n",
    "        | X1 = butler\n",
    "        | X1 = charles ) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_3)).\n",
    "\n",
    "fof(pel55_1,axiom,(\n",
    "    ? [X1] :\n",
    "      ( lives(X1)\n",
    "      & killed(X1,agatha) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_1)).\n",
    "\n",
    "fof(pel55_4,axiom,(\n",
    "    ! [X1,X2] :\n",
    "      ( killed(X1,X2)\n",
    "     => hates(X1,X2) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_4)).\n",
    "\n",
    "fof(pel55_9,axiom,(\n",
    "    ! [X1] :\n",
    "      ( hates(agatha,X1)\n",
    "     => hates(butler,X1) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_9)).\n",
    "\n",
    "fof(pel55_7,axiom,(\n",
    "    ! [X1] :\n",
    "      ( X1 != butler\n",
    "     => hates(agatha,X1) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_7)).\n",
    "\n",
    "fof(pel55_5,axiom,(\n",
    "    ! [X1,X2] :\n",
    "      ( killed(X1,X2)\n",
    "     => ~ richer(X1,X2) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_5)).\n",
    "\n",
    "fof(pel55_8,axiom,(\n",
    "    ! [X1] :\n",
    "      ( ~ richer(X1,agatha)\n",
    "     => hates(butler,X1) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_8)).\n",
    "\n",
    "fof(pel55_6,axiom,(\n",
    "    ! [X1] :\n",
    "      ( hates(agatha,X1)\n",
    "     => ~ hates(charles,X1) ) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_6)).\n",
    "\n",
    "fof(pel55_10,axiom,(\n",
    "    ! [X1] :\n",
    "    ? [X2] : ~ hates(X1,X2) ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_10)).\n",
    "\n",
    "fof(pel55_11,axiom,(\n",
    "    agatha != butler ),\n",
    "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_11)).\n",
    "\n",
    "fof(c_0_11,negated_conjecture,(\n",
    "    ~ killed(agatha,agatha) ),\n",
    "    inference(assume_negation,[status(cth)],[pel55])).\n",
    "\n",
    "fof(c_0_12,plain,(\n",
    "    ! [X4] :\n",
    "      ( ~ lives(X4)\n",
    "      | X4 = agatha\n",
    "      | X4 = butler\n",
    "      | X4 = charles ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_3])])).\n",
    "\n",
    "fof(c_0_13,plain,\n",
    "    ( lives(esk1_0)\n",
    "    & killed(esk1_0,agatha) ),\n",
    "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel55_1])])).\n",
    "\n",
    "fof(c_0_14,negated_conjecture,(\n",
    "    ~ killed(agatha,agatha) ),\n",
    "    inference(fof_simplification,[status(thm)],[c_0_11])).\n",
    "\n",
    "cnf(c_0_15,plain,\n",
    "    ( X1 = agatha\n",
    "    | X1 = butler\n",
    "    | X1 = charles\n",
    "    | ~ lives(X1) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_12])).\n",
    "\n",
    "cnf(c_0_16,plain,\n",
    "    ( lives(esk1_0) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_13])).\n",
    "\n",
    "fof(c_0_17,plain,(\n",
    "    ! [X5,X6] :\n",
    "      ( ~ killed(X5,X6)\n",
    "      | hates(X5,X6) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_4])])).\n",
    "\n",
    "cnf(c_0_18,negated_conjecture,\n",
    "    ( ~ killed(agatha,agatha) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_14])).\n",
    "\n",
    "cnf(c_0_19,plain,\n",
    "    ( agatha = esk1_0\n",
    "    | esk1_0 = butler\n",
    "    | esk1_0 = charles ),\n",
    "    inference(spm,[status(thm)],[c_0_15,c_0_16])).\n",
    "\n",
    "fof(c_0_20,plain,(\n",
    "    ! [X12] :\n",
    "      ( ~ hates(agatha,X12)\n",
    "      | hates(butler,X12) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_9])])).\n",
    "\n",
    "fof(c_0_21,plain,(\n",
    "    ! [X10] :\n",
    "      ( X10 = butler\n",
    "      | hates(agatha,X10) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_7])])).\n",
    "\n",
    "fof(c_0_22,plain,(\n",
    "    ! [X7,X8] :\n",
    "      ( ~ killed(X7,X8)\n",
    "      | ~ richer(X7,X8) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_5])])])).\n",
    "\n",
    "fof(c_0_23,plain,(\n",
    "    ! [X11] :\n",
    "      ( richer(X11,agatha)\n",
    "      | hates(butler,X11) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_8])])])).\n",
    "\n",
    "fof(c_0_24,plain,(\n",
    "    ! [X9] :\n",
    "      ( ~ hates(agatha,X9)\n",
    "      | ~ hates(charles,X9) ) ),\n",
    "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_6])])])).\n",
    "\n",
    "cnf(c_0_25,plain,\n",
    "    ( hates(X1,X2)\n",
    "    | ~ killed(X1,X2) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_17])).\n",
    "\n",
    "cnf(c_0_26,plain,\n",
    "    ( killed(esk1_0,agatha) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_13])).\n",
    "\n",
    "cnf(c_0_27,negated_conjecture,\n",
    "    ( esk1_0 = charles\n",
    "    | esk1_0 = butler\n",
    "    | ~ killed(esk1_0,esk1_0) ),\n",
    "    inference(spm,[status(thm)],[c_0_18,c_0_19])).\n",
    "\n",
    "fof(c_0_28,plain,(\n",
    "    ! [X13] : ~ hates(X13,esk2_1(X13)) ),\n",
    "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_10])])])).\n",
    "\n",
    "cnf(c_0_29,plain,\n",
    "    ( hates(butler,X1)\n",
    "    | ~ hates(agatha,X1) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_20])).\n",
    "\n",
    "cnf(c_0_30,plain,\n",
    "    ( X1 = butler\n",
    "    | hates(agatha,X1) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_21])).\n",
    "\n",
    "cnf(c_0_31,plain,\n",
    "    ( ~ killed(X1,X2)\n",
    "    | ~ richer(X1,X2) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_22])).\n",
    "\n",
    "cnf(c_0_32,plain,\n",
    "    ( richer(X1,agatha)\n",
    "    | hates(butler,X1) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_23])).\n",
    "\n",
    "cnf(c_0_33,plain,\n",
    "    ( ~ hates(agatha,X1)\n",
    "    | ~ hates(charles,X1) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_24])).\n",
    "\n",
    "cnf(c_0_34,plain,\n",
    "    ( hates(esk1_0,agatha) ),\n",
    "    inference(spm,[status(thm)],[c_0_25,c_0_26])).\n",
    "\n",
    "cnf(c_0_35,plain,\n",
    "    ( esk1_0 = charles\n",
    "    | esk1_0 = butler ),\n",
    "    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26,c_0_19]),c_0_27])).\n",
    "\n",
    "cnf(c_0_36,plain,\n",
    "    ( ~ hates(X1,esk2_1(X1)) ),\n",
    "    inference(split_conjunct,[status(thm)],[c_0_28])).\n",
    "\n",
    "cnf(c_0_37,plain,\n",
    "    ( X1 = butler\n",
    "    | hates(butler,X1) ),\n",
    "    inference(spm,[status(thm)],[c_0_29,c_0_30])).\n",
    "\n",
    "cnf(c_0_38,plain,\n",
    "    ( hates(butler,X1)\n",
    "    | ~ killed(X1,agatha) ),\n",
    "    inference(spm,[status(thm)],[c_0_31,c_0_32])).\n",
    "\n",
    "cnf(c_0_39,plain,\n",
    "    ( X1 = butler\n",
    "    | ~ hates(charles,X1) ),\n",
    "    inference(spm,[status(thm)],[c_0_33,c_0_30])).\n",
    "\n",
    "cnf(c_0_40,plain,\n",
    "    ( esk1_0 = butler\n",
    "    | hates(charles,agatha) ),\n",
    "    inference(spm,[status(thm)],[c_0_34,c_0_35])).\n",
    "\n",
    "cnf(c_0_41,plain,\n",
    "    ( agatha != butler ),\n",
    "    inference(split_conjunct,[status(thm)],[pel55_11])).\n",
    "\n",
    "cnf(c_0_42,plain,\n",
    "    ( esk2_1(butler) = butler ),\n",
    "    inference(spm,[status(thm)],[c_0_36,c_0_37])).\n",
    "\n",
    "cnf(c_0_43,plain,\n",
    "    ( hates(butler,esk1_0) ),\n",
    "    inference(spm,[status(thm)],[c_0_38,c_0_26])).\n",
    "\n",
    "cnf(c_0_44,plain,\n",
    "    ( esk1_0 = butler ),\n",
    "    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39,c_0_40]),c_0_41])).\n",
    "\n",
    "cnf(c_0_45,plain,\n",
    "    ( ~ hates(butler,butler) ),\n",
    "    inference(spm,[status(thm)],[c_0_36,c_0_42])).\n",
    "\n",
    "cnf(c_0_46,plain,\n",
    "    ( $false ),\n",
    "    inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_43,c_0_44]),c_0_45]),\n",
    "    [proof]).\n",
    "%------------------------------------------------------------------------------\n",
    "%----ORIGINAL SYSTEM OUTPUT\n",
    "% 0.00/0.04  % Problem    : PUZ001+1 : TPTP v7.3.0. Released v2.0.0.\n",
    "% 0.00/0.05  % Command    : eprover --delete-bad-limit=2000000000 --definitional-cnf=24 -s --print-statistics -R --print-version --proof-object --auto-schedule --cpu-limit=%d %s\n",
    "% 0.03/0.26  % Computer   : n120.star.cs.uiowa.edu\n",
    "% 0.03/0.26  % Model      : x86_64 x86_64\n",
    "% 0.03/0.26  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz\n",
    "% 0.03/0.26  % Memory     : 32218.5MB\n",
    "% 0.03/0.26  % OS         : Linux 3.10.0-862.11.6.el7.x86_64\n",
    "% 0.03/0.26  % CPULimit   : 300\n",
    "% 0.03/0.26  % DateTime   : Tue Jul  9 11:22:36 CDT 2019\n",
    "% 0.03/0.26  % CPUTime    : \n",
    "% 0.03/0.26  # Version: 2.3_dev009\n",
    "% 0.03/0.26  # No SInE strategy applied\n",
    "% 0.03/0.26  # Trying AutoSched0 for 151 seconds\n",
    "% 0.03/0.27  # AutoSched0-Mode selected heuristic G_E___200_C45_F1_SE_CS_SP_PI_CO_S0V\n",
    "% 0.03/0.27  # and selection function PSelectComplexExceptRRHorn.\n",
    "% 0.03/0.27  #\n",
    "% 0.03/0.27  # Preprocessing time       : 0.012 s\n",
    "% 0.03/0.27  \n",
    "% 0.03/0.27  # Proof found!\n",
    "% 0.03/0.27  # SZS status Theorem\n",
    "% 0.03/0.27  # SZS output start CNFRefutation\n",
    "% 0.03/0.27  fof(pel55, conjecture, killed(agatha,agatha), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55)).\n",
    "% 0.03/0.27  fof(pel55_3, axiom, ![X1]:(lives(X1)=>(X1=agatha|X1=butler|X1=charles)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_3)).\n",
    "% 0.03/0.27  fof(pel55_1, axiom, ?[X1]:(lives(X1)&killed(X1,agatha)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_1)).\n",
    "% 0.03/0.27  fof(pel55_4, axiom, ![X1, X2]:(killed(X1,X2)=>hates(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_4)).\n",
    "% 0.03/0.27  fof(pel55_9, axiom, ![X1]:(hates(agatha,X1)=>hates(butler,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_9)).\n",
    "% 0.03/0.27  fof(pel55_7, axiom, ![X1]:(X1!=butler=>hates(agatha,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_7)).\n",
    "% 0.03/0.27  fof(pel55_5, axiom, ![X1, X2]:(killed(X1,X2)=>~(richer(X1,X2))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_5)).\n",
    "% 0.03/0.27  fof(pel55_8, axiom, ![X1]:(~(richer(X1,agatha))=>hates(butler,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_8)).\n",
    "% 0.03/0.27  fof(pel55_6, axiom, ![X1]:(hates(agatha,X1)=>~(hates(charles,X1))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_6)).\n",
    "% 0.03/0.27  fof(pel55_10, axiom, ![X1]:?[X2]:~(hates(X1,X2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_10)).\n",
    "% 0.03/0.27  fof(pel55_11, axiom, agatha!=butler, file('/export/starexec/sandbox/benchmark/theBenchmark.p', pel55_11)).\n",
    "% 0.03/0.27  fof(c_0_11, negated_conjecture, ~(killed(agatha,agatha)), inference(assume_negation,[status(cth)],[pel55])).\n",
    "% 0.03/0.27  fof(c_0_12, plain, ![X4]:(~lives(X4)|(X4=agatha|X4=butler|X4=charles)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_3])])).\n",
    "% 0.03/0.27  fof(c_0_13, plain, (lives(esk1_0)&killed(esk1_0,agatha)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel55_1])])).\n",
    "% 0.03/0.27  fof(c_0_14, negated_conjecture, ~killed(agatha,agatha), inference(fof_simplification,[status(thm)],[c_0_11])).\n",
    "% 0.03/0.27  cnf(c_0_15, plain, (X1=agatha|X1=butler|X1=charles|~lives(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).\n",
    "% 0.03/0.27  cnf(c_0_16, plain, (lives(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).\n",
    "% 0.03/0.27  fof(c_0_17, plain, ![X5, X6]:(~killed(X5,X6)|hates(X5,X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_4])])).\n",
    "% 0.03/0.27  cnf(c_0_18, negated_conjecture, (~killed(agatha,agatha)), inference(split_conjunct,[status(thm)],[c_0_14])).\n",
    "% 0.03/0.27  cnf(c_0_19, plain, (agatha=esk1_0|esk1_0=butler|esk1_0=charles), inference(spm,[status(thm)],[c_0_15, c_0_16])).\n",
    "% 0.03/0.27  fof(c_0_20, plain, ![X12]:(~hates(agatha,X12)|hates(butler,X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_9])])).\n",
    "% 0.03/0.27  fof(c_0_21, plain, ![X10]:(X10=butler|hates(agatha,X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_7])])).\n",
    "% 0.03/0.27  fof(c_0_22, plain, ![X7, X8]:(~killed(X7,X8)|~richer(X7,X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_5])])])).\n",
    "% 0.03/0.27  fof(c_0_23, plain, ![X11]:(richer(X11,agatha)|hates(butler,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_8])])])).\n",
    "% 0.03/0.27  fof(c_0_24, plain, ![X9]:(~hates(agatha,X9)|~hates(charles,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_6])])])).\n",
    "% 0.03/0.27  cnf(c_0_25, plain, (hates(X1,X2)|~killed(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).\n",
    "% 0.03/0.27  cnf(c_0_26, plain, (killed(esk1_0,agatha)), inference(split_conjunct,[status(thm)],[c_0_13])).\n",
    "% 0.03/0.27  cnf(c_0_27, negated_conjecture, (esk1_0=charles|esk1_0=butler|~killed(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).\n",
    "% 0.03/0.27  fof(c_0_28, plain, ![X13]:~hates(X13,esk2_1(X13)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_10])])])).\n",
    "% 0.03/0.27  cnf(c_0_29, plain, (hates(butler,X1)|~hates(agatha,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).\n",
    "% 0.03/0.27  cnf(c_0_30, plain, (X1=butler|hates(agatha,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).\n",
    "% 0.03/0.27  cnf(c_0_31, plain, (~killed(X1,X2)|~richer(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).\n",
    "% 0.03/0.27  cnf(c_0_32, plain, (richer(X1,agatha)|hates(butler,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).\n",
    "% 0.03/0.27  cnf(c_0_33, plain, (~hates(agatha,X1)|~hates(charles,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).\n",
    "% 0.03/0.27  cnf(c_0_34, plain, (hates(esk1_0,agatha)), inference(spm,[status(thm)],[c_0_25, c_0_26])).\n",
    "% 0.03/0.27  cnf(c_0_35, plain, (esk1_0=charles|esk1_0=butler), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_19]), c_0_27])).\n",
    "% 0.03/0.27  cnf(c_0_36, plain, (~hates(X1,esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).\n",
    "% 0.03/0.27  cnf(c_0_37, plain, (X1=butler|hates(butler,X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).\n",
    "% 0.03/0.27  cnf(c_0_38, plain, (hates(butler,X1)|~killed(X1,agatha)), inference(spm,[status(thm)],[c_0_31, c_0_32])).\n",
    "% 0.03/0.27  cnf(c_0_39, plain, (X1=butler|~hates(charles,X1)), inference(spm,[status(thm)],[c_0_33, c_0_30])).\n",
    "% 0.03/0.27  cnf(c_0_40, plain, (esk1_0=butler|hates(charles,agatha)), inference(spm,[status(thm)],[c_0_34, c_0_35])).\n",
    "% 0.03/0.27  cnf(c_0_41, plain, (agatha!=butler), inference(split_conjunct,[status(thm)],[pel55_11])).\n",
    "% 0.03/0.27  cnf(c_0_42, plain, (esk2_1(butler)=butler), inference(spm,[status(thm)],[c_0_36, c_0_37])).\n",
    "% 0.03/0.27  cnf(c_0_43, plain, (hates(butler,esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_26])).\n",
    "% 0.03/0.27  cnf(c_0_44, plain, (esk1_0=butler), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).\n",
    "% 0.03/0.27  cnf(c_0_45, plain, (~hates(butler,butler)), inference(spm,[status(thm)],[c_0_36, c_0_42])).\n",
    "% 0.03/0.27  cnf(c_0_46, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45]), ['proof']).\n",
    "% 0.03/0.27  # SZS output end CNFRefutation\n",
    "% 0.03/0.27  # Proof object total steps             : 47\n",
    "% 0.03/0.27  # Proof object clause steps            : 25\n",
    "% 0.03/0.27  # Proof object formula steps           : 22\n",
    "% 0.03/0.27  # Proof object conjectures             : 5\n",
    "% 0.03/0.27  # Proof object clause conjectures      : 2\n",
    "% 0.03/0.27  # Proof object formula conjectures     : 3\n",
    "% 0.03/0.27  # Proof object initial clauses used    : 12\n",
    "% 0.03/0.27  # Proof object initial formulas used   : 11\n",
    "% 0.03/0.27  # Proof object generating inferences   : 12\n",
    "% 0.03/0.27  # Proof object simplifying inferences  : 4\n",
    "% 0.03/0.27  # Training examples: 0 positive, 0 negative\n",
    "% 0.03/0.27  # Parsed axioms                        : 14\n",
    "% 0.03/0.27  # Removed by relevancy pruning/SinE    : 0\n",
    "% 0.03/0.27  # Initial clauses                      : 15\n",
    "% 0.03/0.27  # Removed in clause preprocessing      : 0\n",
    "% 0.03/0.27  # Initial clauses in saturation        : 15\n",
    "% 0.03/0.27  # Processed clauses                    : 36\n",
    "% 0.03/0.27  # ...of these trivial                  : 0\n",
    "% 0.03/0.27  # ...subsumed                          : 3\n",
    "% 0.03/0.27  # ...remaining for further processing  : 33\n",
    "% 0.03/0.27  # Other redundant clauses eliminated   : 0\n",
    "% 0.03/0.27  # Clauses deleted for lack of memory   : 0\n",
    "% 0.03/0.27  # Backward-subsumed                    : 3\n",
    "% 0.03/0.27  # Backward-rewritten                   : 7\n",
    "% 0.03/0.27  # Generated clauses                    : 40\n",
    "% 0.03/0.27  # ...of the previous two non-trivial   : 33\n",
    "% 0.03/0.27  # Contextual simplify-reflections      : 1\n",
    "% 0.03/0.27  # Paramodulations                      : 35\n",
    "% 0.03/0.27  # Factorizations                       : 5\n",
    "% 0.03/0.27  # Equation resolutions                 : 0\n",
    "% 0.03/0.27  # Propositional unsat checks           : 0\n",
    "% 0.03/0.27  #    Propositional check models        : 0\n",
    "% 0.03/0.27  #    Propositional check unsatisfiable : 0\n",
    "% 0.03/0.27  #    Propositional clauses             : 0\n",
    "% 0.03/0.27  #    Propositional clauses after purity: 0\n",
    "% 0.03/0.27  #    Propositional unsat core size     : 0\n",
    "% 0.03/0.27  #    Propositional preprocessing time  : 0.000\n",
    "% 0.03/0.27  #    Propositional encoding time       : 0.000\n",
    "% 0.03/0.27  #    Propositional solver time         : 0.000\n",
    "% 0.03/0.27  #    Success case prop preproc time    : 0.000\n",
    "% 0.03/0.27  #    Success case prop encoding time   : 0.000\n",
    "% 0.03/0.27  #    Success case prop solver time     : 0.000\n",
    "% 0.03/0.27  # Current number of processed clauses  : 23\n",
    "% 0.03/0.27  #    Positive orientable unit clauses  : 6\n",
    "% 0.03/0.27  #    Positive unorientable unit clauses: 0\n",
    "% 0.03/0.27  #    Negative unit clauses             : 6\n",
    "% 0.03/0.27  #    Non-unit-clauses                  : 11\n",
    "% 0.03/0.27  # Current number of unprocessed clauses: 11\n",
    "% 0.03/0.27  # ...number of literals in the above   : 38\n",
    "% 0.03/0.27  # Current number of archived formulas  : 0\n",
    "% 0.03/0.27  # Current number of archived clauses   : 10\n",
    "% 0.03/0.27  # Clause-clause subsumption calls (NU) : 12\n",
    "% 0.03/0.27  # Rec. Clause-clause subsumption calls : 10\n",
    "% 0.03/0.27  # Non-unit clause-clause subsumptions  : 3\n",
    "% 0.03/0.27  # Unit Clause-clause subsumption calls : 7\n",
    "% 0.03/0.27  # Rewrite failures with RHS unbound    : 0\n",
    "% 0.03/0.27  # BW rewrite match attempts            : 1\n",
    "% 0.03/0.27  # BW rewrite match successes           : 1\n",
    "% 0.03/0.27  # Condensation attempts                : 36\n",
    "% 0.03/0.27  # Condensation successes               : 0\n",
    "% 0.03/0.27  # Termbank termtop insertions          : 887\n",
    "% 0.03/0.27  \n",
    "% 0.03/0.27  # -------------------------------------------------\n",
    "% 0.03/0.27  # User time                : 0.012 s\n",
    "% 0.03/0.27  # System time              : 0.004 s\n",
    "% 0.03/0.27  # Total time               : 0.015 s\n",
    "% 0.03/0.27  # Maximum resident set size: 1564 pages\n",
    "%------------------------------------------------------------------------------\n",
    "\"\"\""
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 22,
   "metadata": {},
   "outputs": [],
   "source": [
    "args = {\n",
    "    \"ProblemSource\": \"FORMULAE\",\n",
    "    \"FORMULAEProblem\": problem,\n",
    "    \"SolutionFormat\": \"TPTP\",\n",
    "    \n",
    "    \"QuietFlag\": \"-q01\",\n",
    "    \"SubmitButton\": \"ProcessSolution\",\n",
    "    \n",
    "    \n",
    "    \"System___AGInTRater---0.0\": \"AGInTRater---0.0\",\n",
    "    \"TimeLimit___AGInTRater---0.0\": \"60\",\n",
    "    \"Transform___AGInTRater---0.0\": \"none\",\n",
    "    \"Format___AGInTRater---0.0\": \"tptp:raw\",\n",
    "    \"Command___AGInTRater---0.0\": \"AGInTRater -c %s\"\n",
    "}"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 18,
   "metadata": {},
   "outputs": [],
   "source": [
    "requests.post?"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 23,
   "metadata": {},
   "outputs": [],
   "source": [
    "response = requests.post(url, args)\n",
    "text = response.text"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 24,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "<!DOCTYPE html\n",
      "\tPUBLIC \"-//W3C//DTD XHTML 1.0 Transitional//EN\"\n",
      "\t \"http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd\">\n",
      "<html xmlns=\"http://www.w3.org/1999/xhtml\" lang=\"en-US\" xml:lang=\"en-US\">\n",
      "<head>\n",
      "<title>System Results</title>\n",
      "<meta name=\"ROBOTS\" content=\"INDEX,NOFOLLOW\" />\n",
      "<script type=\"text/javascript\">//<![CDATA[\n",
      "\n",
      "  var tstp_dump;\n",
      "  function openSoTSTP (dump) {   \n",
      "    var tstp_url = 'http://tptp.org/cgi-bin/SystemOnTSTP';\n",
      "    var tstp_browser = window.open(tstp_url, '_blank');\n",
      "    tstp_dump = dump;\n",
      "  }\n",
      "  function getTSTPDump () {\n",
      "    return tstp_dump;\n",
      "  }\n",
      "\n",
      "//]]></script>\n",
      "<meta http-equiv=\"Content-Type\" content=\"text/html; charset=iso-8859-1\" />\n",
      "</head>\n",
      "<body>\n",
      "\n",
      "% SZS start RequiredInformation<BR>\n",
      "% Congratulations - you have become a registered power user of SystemOnTPTP, at IP address 137.103.55.74.<BR>\n",
      "% Please consider donating to the TPTP project - see tptp.org for details.<BR>\n",
      "% When you donate this message will disappear.<BR>\n",
      "% If you do not donate a random delay might be added to your processing time.<BR>\n",
      "% SZS end RequiredInformation<BR>\n",
      "<PRE>\n",
      "% START OF SYSTEM OUTPUT\n",
      "% SUCCESS: Rated\n",
      "% AGInTRater output start\n",
      "cnf(c_0_45,plain,\n",
      "    ( ~ hates(butler,butler) ),\n",
      "    inference(spm,[status(thm)],[c_0_36,c_0_42]),\n",
      "    [interesting(1.00),obviousness(1.00),complexity(1.00),intensity(0.65),surprisingness(1.00),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_37,plain,\n",
      "    ( X1 = butler\n",
      "    | hates(butler,X1) ),\n",
      "    inference(spm,[status(thm)],[c_0_29,c_0_30]),\n",
      "    [interesting(0.96),obviousness(0.00),complexity(0.67),intensity(0.87),surprisingness(0.81),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_44,plain,\n",
      "    ( esk1_0 = butler ),\n",
      "    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39,c_0_40]),c_0_41]),\n",
      "    [interesting(0.96),obviousness(1.00),complexity(0.75),intensity(1.00),surprisingness(0.84),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(0.00)]).\n",
      "\n",
      "cnf(c_0_42,plain,\n",
      "    ( esk2_1(butler) = butler ),\n",
      "    inference(spm,[status(thm)],[c_0_36,c_0_37]),\n",
      "    [interesting(0.85),obviousness(0.50),complexity(0.67),intensity(0.87),surprisingness(0.74),adaptivity(0.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_43,plain,\n",
      "    ( hates(butler,esk1_0) ),\n",
      "    inference(spm,[status(thm)],[c_0_38,c_0_26]),\n",
      "    [interesting(0.81),obviousness(0.50),complexity(0.67),intensity(0.87),surprisingness(0.71),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_39,plain,\n",
      "    ( X1 = butler\n",
      "    | ~ hates(charles,X1) ),\n",
      "    inference(spm,[status(thm)],[c_0_33,c_0_30]),\n",
      "    [interesting(0.78),obviousness(0.00),complexity(0.33),intensity(1.00),surprisingness(0.45),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_38,plain,\n",
      "    ( hates(butler,X1)\n",
      "    | ~ killed(X1,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_31,c_0_32]),\n",
      "    [interesting(0.76),obviousness(0.00),complexity(0.33),intensity(0.48),surprisingness(0.47),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_34,plain,\n",
      "    ( hates(esk1_0,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_25,c_0_26]),\n",
      "    [interesting(0.63),obviousness(0.00),complexity(0.67),intensity(0.55),surprisingness(0.73),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_35,plain,\n",
      "    ( esk1_0 = charles\n",
      "    | esk1_0 = butler ),\n",
      "    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26,c_0_19]),c_0_27]),\n",
      "    [interesting(0.53),obviousness(0.43),complexity(0.50),intensity(0.22),surprisingness(0.60),adaptivity(0.00),focus(ignored),weight(0.50),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_40,plain,\n",
      "    ( esk1_0 = butler\n",
      "    | hates(charles,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_34,c_0_35]),\n",
      "    [interesting(0.35),obviousness(0.71),complexity(0.00),intensity(0.48),surprisingness(0.00),adaptivity(0.00),focus(ignored),weight(0.50),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_27,negated_conjecture,\n",
      "    ( esk1_0 = charles\n",
      "    | esk1_0 = butler\n",
      "    | ~ killed(esk1_0,esk1_0) ),\n",
      "    inference(spm,[status(thm)],[c_0_18,c_0_19]),\n",
      "    [interesting(0.03),obviousness(0.14),complexity(0.25),intensity(0.04),surprisingness(0.26),adaptivity(0.00),focus(ignored),weight(0.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_19,plain,\n",
      "    ( agatha = esk1_0\n",
      "    | esk1_0 = butler\n",
      "    | esk1_0 = charles ),\n",
      "    inference(spm,[status(thm)],[c_0_15,c_0_16]),\n",
      "    [interesting(0.00),obviousness(0.00),complexity(0.00),intensity(0.00),surprisingness(0.00),adaptivity(0.00),focus(ignored),weight(0.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_46,plain,\n",
      "    ( $false ),\n",
      "    inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_43,c_0_44]),c_0_45]),\n",
      "    [proof,interesting(0.00)]).\n",
      "\n",
      "cnf(c_0_31,plain,\n",
      "    ( ~ killed(X1,X2)\n",
      "    | ~ richer(X1,X2) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_22]),\n",
      "    [interesting(-0.03),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.98),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_25,plain,\n",
      "    ( hates(X1,X2)\n",
      "    | ~ killed(X1,X2) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_17]),\n",
      "    [interesting(-0.04),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.96),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_36,plain,\n",
      "    ( ~ hates(X1,esk2_1(X1)) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_28]),\n",
      "    [interesting(-0.04),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_29,plain,\n",
      "    ( hates(butler,X1)\n",
      "    | ~ hates(agatha,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_20]),\n",
      "    [interesting(-0.19),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.86),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_33,plain,\n",
      "    ( ~ hates(agatha,X1)\n",
      "    | ~ hates(charles,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_24]),\n",
      "    [interesting(-0.21),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.75),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_16,plain,\n",
      "    ( lives(esk1_0) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_13]),\n",
      "    [interesting(-0.25),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_17,plain,(\n",
      "    ! [X5,X6] :\n",
      "      ( ~ killed(X5,X6)\n",
      "      | hates(X5,X6) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_4])]),\n",
      "    [interesting(-0.27),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.96),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_22,plain,(\n",
      "    ! [X7,X8] :\n",
      "      ( ~ killed(X7,X8)\n",
      "      | ~ richer(X7,X8) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_5])])]),\n",
      "    [interesting(-0.27),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_28,plain,(\n",
      "    ! [X13] : ~ hates(X13,esk2_1(X13)) ),\n",
      "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_10])])]),\n",
      "    [interesting(-0.28),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_18,negated_conjecture,\n",
      "    ( ~ killed(agatha,agatha) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_14]),\n",
      "    [interesting(-0.31),obviousness(0.07),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_14,negated_conjecture,(\n",
      "    ~ killed(agatha,agatha) ),\n",
      "    inference(fof_simplification,[status(thm)],[c_0_11]),\n",
      "    [interesting(-0.32),obviousness(0.03),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_11,negated_conjecture,(\n",
      "    ~ killed(agatha,agatha) ),\n",
      "    inference(assume_negation,[status(cth)],[pel55]),\n",
      "    [interesting(-0.33),obviousness(0.00),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_30,plain,\n",
      "    ( X1 = butler\n",
      "    | hates(agatha,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_21]),\n",
      "    [interesting(-0.33),obviousness(0.10),complexity(0.33),intensity(0.22),surprisingness(0.55),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_32,plain,\n",
      "    ( richer(X1,agatha)\n",
      "    | hates(butler,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_23]),\n",
      "    [interesting(-0.35),obviousness(0.10),complexity(0.33),intensity(0.22),surprisingness(0.42),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_20,plain,(\n",
      "    ! [X12] :\n",
      "      ( ~ hates(agatha,X12)\n",
      "      | hates(butler,X12) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_9])]),\n",
      "    [interesting(-0.42),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.86),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_26,plain,\n",
      "    ( killed(esk1_0,agatha) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_13]),\n",
      "    [interesting(-0.42),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.62),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_24,plain,(\n",
      "    ! [X9] :\n",
      "      ( ~ hates(agatha,X9)\n",
      "      | ~ hates(charles,X9) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_6])])]),\n",
      "    [interesting(-0.45),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.75),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_41,plain,\n",
      "    ( agatha != butler ),\n",
      "    inference(split_conjunct,[status(thm)],[pel55_11]),\n",
      "    [interesting(-0.49),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.89),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(ignored)]).\n",
      "\n",
      "fof(c_0_21,plain,(\n",
      "    ! [X10] :\n",
      "      ( X10 = butler\n",
      "      | hates(agatha,X10) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_7])]),\n",
      "    [interesting(-0.56),obviousness(0.00),complexity(0.33),intensity(0.22),surprisingness(0.55),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_23,plain,(\n",
      "    ! [X11] :\n",
      "      ( richer(X11,agatha)\n",
      "      | hates(butler,X11) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_8])])]),\n",
      "    [interesting(-0.59),obviousness(0.00),complexity(0.33),intensity(0.22),surprisingness(0.42),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_13,plain,\n",
      "    ( lives(esk1_0)\n",
      "    & killed(esk1_0,agatha) ),\n",
      "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel55_1])]),\n",
      "    [interesting(-0.71),obviousness(0.00),complexity(0.33),intensity(0.00),surprisingness(0.20),adaptivity(0.00),focus(1.00),weight(0.57),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_12,plain,(\n",
      "    ! [X4] :\n",
      "      ( ~ lives(X4)\n",
      "      | X4 = agatha\n",
      "      | X4 = butler\n",
      "      | X4 = charles ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_3])]),\n",
      "    [interesting(-0.84),obviousness(0.00),complexity(0.00),intensity(0.22),surprisingness(0.01),adaptivity(0.00),focus(1.00),weight(0.29),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_15,plain,\n",
      "    ( X1 = agatha\n",
      "    | X1 = butler\n",
      "    | X1 = charles\n",
      "    | ~ lives(X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_12]),\n",
      "    [interesting(-0.93),obviousness(0.10),complexity(0.00),intensity(0.22),surprisingness(0.01),adaptivity(0.50),focus(0.00),weight(0.29),usefulness(1.00)]).\n",
      "\n",
      "fof(pel55,conjecture,(\n",
      "    killed(agatha,agatha) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_3,axiom,(\n",
      "    ! [X1] :\n",
      "      ( lives(X1)\n",
      "     =&gt; ( X1 = agatha\n",
      "        | X1 = butler\n",
      "        | X1 = charles ) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_3),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_1,axiom,(\n",
      "    ? [X1] :\n",
      "      ( lives(X1)\n",
      "      & killed(X1,agatha) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_1),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_4,axiom,(\n",
      "    ! [X1,X2] :\n",
      "      ( killed(X1,X2)\n",
      "     =&gt; hates(X1,X2) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_4),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_9,axiom,(\n",
      "    ! [X1] :\n",
      "      ( hates(agatha,X1)\n",
      "     =&gt; hates(butler,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_9),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_7,axiom,(\n",
      "    ! [X1] :\n",
      "      ( X1 != butler\n",
      "     =&gt; hates(agatha,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_7),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_5,axiom,(\n",
      "    ! [X1,X2] :\n",
      "      ( killed(X1,X2)\n",
      "     =&gt; ~ richer(X1,X2) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_5),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_8,axiom,(\n",
      "    ! [X1] :\n",
      "      ( ~ richer(X1,agatha)\n",
      "     =&gt; hates(butler,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_8),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_6,axiom,(\n",
      "    ! [X1] :\n",
      "      ( hates(agatha,X1)\n",
      "     =&gt; ~ hates(charles,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_6),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_10,axiom,(\n",
      "    ! [X1] :\n",
      "    ? [X2] : ~ hates(X1,X2) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_10),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_11,axiom,(\n",
      "    agatha != butler ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_11),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "% AGInTRater output end\n",
      "\n",
      "% END OF SYSTEM OUTPUT\n",
      "% RESULT: SOT_nr4rgR - AGInTRater---0.0 says Rated - CPU = 0.00 WC = 0.09 \n",
      "% OUTPUT: SOT_nr4rgR - AGInTRater---0.0 says Derivation - CPU = 0.00 WC = 0.09 \n",
      "</PRE>\n",
      "\n",
      "</body>\n",
      "</html>\n"
     ]
    }
   ],
   "source": [
    "print(text)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 25,
   "metadata": {},
   "outputs": [],
   "source": [
    "begin = text.index(\"<PRE>\") + 5\n",
    "end = text.index(\"</PRE>\")"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 26,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "\n",
      "% START OF SYSTEM OUTPUT\n",
      "% SUCCESS: Rated\n",
      "% AGInTRater output start\n",
      "cnf(c_0_45,plain,\n",
      "    ( ~ hates(butler,butler) ),\n",
      "    inference(spm,[status(thm)],[c_0_36,c_0_42]),\n",
      "    [interesting(1.00),obviousness(1.00),complexity(1.00),intensity(0.65),surprisingness(1.00),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_37,plain,\n",
      "    ( X1 = butler\n",
      "    | hates(butler,X1) ),\n",
      "    inference(spm,[status(thm)],[c_0_29,c_0_30]),\n",
      "    [interesting(0.96),obviousness(0.00),complexity(0.67),intensity(0.87),surprisingness(0.81),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_44,plain,\n",
      "    ( esk1_0 = butler ),\n",
      "    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39,c_0_40]),c_0_41]),\n",
      "    [interesting(0.96),obviousness(1.00),complexity(0.75),intensity(1.00),surprisingness(0.84),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(0.00)]).\n",
      "\n",
      "cnf(c_0_42,plain,\n",
      "    ( esk2_1(butler) = butler ),\n",
      "    inference(spm,[status(thm)],[c_0_36,c_0_37]),\n",
      "    [interesting(0.85),obviousness(0.50),complexity(0.67),intensity(0.87),surprisingness(0.74),adaptivity(0.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_43,plain,\n",
      "    ( hates(butler,esk1_0) ),\n",
      "    inference(spm,[status(thm)],[c_0_38,c_0_26]),\n",
      "    [interesting(0.81),obviousness(0.50),complexity(0.67),intensity(0.87),surprisingness(0.71),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_39,plain,\n",
      "    ( X1 = butler\n",
      "    | ~ hates(charles,X1) ),\n",
      "    inference(spm,[status(thm)],[c_0_33,c_0_30]),\n",
      "    [interesting(0.78),obviousness(0.00),complexity(0.33),intensity(1.00),surprisingness(0.45),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_38,plain,\n",
      "    ( hates(butler,X1)\n",
      "    | ~ killed(X1,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_31,c_0_32]),\n",
      "    [interesting(0.76),obviousness(0.00),complexity(0.33),intensity(0.48),surprisingness(0.47),adaptivity(1.00),focus(ignored),weight(0.83),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_34,plain,\n",
      "    ( hates(esk1_0,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_25,c_0_26]),\n",
      "    [interesting(0.63),obviousness(0.00),complexity(0.67),intensity(0.55),surprisingness(0.73),adaptivity(0.00),focus(ignored),weight(1.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_35,plain,\n",
      "    ( esk1_0 = charles\n",
      "    | esk1_0 = butler ),\n",
      "    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26,c_0_19]),c_0_27]),\n",
      "    [interesting(0.53),obviousness(0.43),complexity(0.50),intensity(0.22),surprisingness(0.60),adaptivity(0.00),focus(ignored),weight(0.50),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_40,plain,\n",
      "    ( esk1_0 = butler\n",
      "    | hates(charles,agatha) ),\n",
      "    inference(spm,[status(thm)],[c_0_34,c_0_35]),\n",
      "    [interesting(0.35),obviousness(0.71),complexity(0.00),intensity(0.48),surprisingness(0.00),adaptivity(0.00),focus(ignored),weight(0.50),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_27,negated_conjecture,\n",
      "    ( esk1_0 = charles\n",
      "    | esk1_0 = butler\n",
      "    | ~ killed(esk1_0,esk1_0) ),\n",
      "    inference(spm,[status(thm)],[c_0_18,c_0_19]),\n",
      "    [interesting(0.03),obviousness(0.14),complexity(0.25),intensity(0.04),surprisingness(0.26),adaptivity(0.00),focus(ignored),weight(0.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_19,plain,\n",
      "    ( agatha = esk1_0\n",
      "    | esk1_0 = butler\n",
      "    | esk1_0 = charles ),\n",
      "    inference(spm,[status(thm)],[c_0_15,c_0_16]),\n",
      "    [interesting(0.00),obviousness(0.00),complexity(0.00),intensity(0.00),surprisingness(0.00),adaptivity(0.00),focus(ignored),weight(0.00),usefulness(ignored)]).\n",
      "\n",
      "cnf(c_0_46,plain,\n",
      "    ( $false ),\n",
      "    inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_43,c_0_44]),c_0_45]),\n",
      "    [proof,interesting(0.00)]).\n",
      "\n",
      "cnf(c_0_31,plain,\n",
      "    ( ~ killed(X1,X2)\n",
      "    | ~ richer(X1,X2) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_22]),\n",
      "    [interesting(-0.03),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.98),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_25,plain,\n",
      "    ( hates(X1,X2)\n",
      "    | ~ killed(X1,X2) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_17]),\n",
      "    [interesting(-0.04),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.96),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_36,plain,\n",
      "    ( ~ hates(X1,esk2_1(X1)) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_28]),\n",
      "    [interesting(-0.04),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(1.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_29,plain,\n",
      "    ( hates(butler,X1)\n",
      "    | ~ hates(agatha,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_20]),\n",
      "    [interesting(-0.19),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.86),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_33,plain,\n",
      "    ( ~ hates(agatha,X1)\n",
      "    | ~ hates(charles,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_24]),\n",
      "    [interesting(-0.21),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.75),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_16,plain,\n",
      "    ( lives(esk1_0) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_13]),\n",
      "    [interesting(-0.25),obviousness(0.10),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_17,plain,(\n",
      "    ! [X5,X6] :\n",
      "      ( ~ killed(X5,X6)\n",
      "      | hates(X5,X6) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_4])]),\n",
      "    [interesting(-0.27),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.96),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_22,plain,(\n",
      "    ! [X7,X8] :\n",
      "      ( ~ killed(X7,X8)\n",
      "      | ~ richer(X7,X8) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_5])])]),\n",
      "    [interesting(-0.27),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_28,plain,(\n",
      "    ! [X13] : ~ hates(X13,esk2_1(X13)) ),\n",
      "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_10])])]),\n",
      "    [interesting(-0.28),obviousness(0.00),complexity(1.00),intensity(0.22),surprisingness(0.94),adaptivity(0.00),focus(1.00),weight(1.00),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_18,negated_conjecture,\n",
      "    ( ~ killed(agatha,agatha) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_14]),\n",
      "    [interesting(-0.31),obviousness(0.07),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_14,negated_conjecture,(\n",
      "    ~ killed(agatha,agatha) ),\n",
      "    inference(fof_simplification,[status(thm)],[c_0_11]),\n",
      "    [interesting(-0.32),obviousness(0.03),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_11,negated_conjecture,(\n",
      "    ~ killed(agatha,agatha) ),\n",
      "    inference(assume_negation,[status(cth)],[pel55]),\n",
      "    [interesting(-0.33),obviousness(0.00),complexity(1.00),intensity(0.07),surprisingness(0.98),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_30,plain,\n",
      "    ( X1 = butler\n",
      "    | hates(agatha,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_21]),\n",
      "    [interesting(-0.33),obviousness(0.10),complexity(0.33),intensity(0.22),surprisingness(0.55),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_32,plain,\n",
      "    ( richer(X1,agatha)\n",
      "    | hates(butler,X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_23]),\n",
      "    [interesting(-0.35),obviousness(0.10),complexity(0.33),intensity(0.22),surprisingness(0.42),adaptivity(1.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_20,plain,(\n",
      "    ! [X12] :\n",
      "      ( ~ hates(agatha,X12)\n",
      "      | hates(butler,X12) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_9])]),\n",
      "    [interesting(-0.42),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.86),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_26,plain,\n",
      "    ( killed(esk1_0,agatha) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_13]),\n",
      "    [interesting(-0.42),obviousness(0.10),complexity(0.67),intensity(0.22),surprisingness(0.62),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_24,plain,(\n",
      "    ! [X9] :\n",
      "      ( ~ hates(agatha,X9)\n",
      "      | ~ hates(charles,X9) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_6])])]),\n",
      "    [interesting(-0.45),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.75),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_41,plain,\n",
      "    ( agatha != butler ),\n",
      "    inference(split_conjunct,[status(thm)],[pel55_11]),\n",
      "    [interesting(-0.49),obviousness(0.00),complexity(0.67),intensity(0.22),surprisingness(0.89),adaptivity(0.00),focus(1.00),weight(0.86),usefulness(ignored)]).\n",
      "\n",
      "fof(c_0_21,plain,(\n",
      "    ! [X10] :\n",
      "      ( X10 = butler\n",
      "      | hates(agatha,X10) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_7])]),\n",
      "    [interesting(-0.56),obviousness(0.00),complexity(0.33),intensity(0.22),surprisingness(0.55),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_23,plain,(\n",
      "    ! [X11] :\n",
      "      ( richer(X11,agatha)\n",
      "      | hates(butler,X11) ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[pel55_8])])]),\n",
      "    [interesting(-0.59),obviousness(0.00),complexity(0.33),intensity(0.22),surprisingness(0.42),adaptivity(0.00),focus(1.00),weight(0.71),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_13,plain,\n",
      "    ( lives(esk1_0)\n",
      "    & killed(esk1_0,agatha) ),\n",
      "    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel55_1])]),\n",
      "    [interesting(-0.71),obviousness(0.00),complexity(0.33),intensity(0.00),surprisingness(0.20),adaptivity(0.00),focus(1.00),weight(0.57),usefulness(1.00)]).\n",
      "\n",
      "fof(c_0_12,plain,(\n",
      "    ! [X4] :\n",
      "      ( ~ lives(X4)\n",
      "      | X4 = agatha\n",
      "      | X4 = butler\n",
      "      | X4 = charles ) ),\n",
      "    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel55_3])]),\n",
      "    [interesting(-0.84),obviousness(0.00),complexity(0.00),intensity(0.22),surprisingness(0.01),adaptivity(0.00),focus(1.00),weight(0.29),usefulness(1.00)]).\n",
      "\n",
      "cnf(c_0_15,plain,\n",
      "    ( X1 = agatha\n",
      "    | X1 = butler\n",
      "    | X1 = charles\n",
      "    | ~ lives(X1) ),\n",
      "    inference(split_conjunct,[status(thm)],[c_0_12]),\n",
      "    [interesting(-0.93),obviousness(0.10),complexity(0.00),intensity(0.22),surprisingness(0.01),adaptivity(0.50),focus(0.00),weight(0.29),usefulness(1.00)]).\n",
      "\n",
      "fof(pel55,conjecture,(\n",
      "    killed(agatha,agatha) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_3,axiom,(\n",
      "    ! [X1] :\n",
      "      ( lives(X1)\n",
      "     =&gt; ( X1 = agatha\n",
      "        | X1 = butler\n",
      "        | X1 = charles ) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_3),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_1,axiom,(\n",
      "    ? [X1] :\n",
      "      ( lives(X1)\n",
      "      & killed(X1,agatha) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_1),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_4,axiom,(\n",
      "    ! [X1,X2] :\n",
      "      ( killed(X1,X2)\n",
      "     =&gt; hates(X1,X2) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_4),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_9,axiom,(\n",
      "    ! [X1] :\n",
      "      ( hates(agatha,X1)\n",
      "     =&gt; hates(butler,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_9),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_7,axiom,(\n",
      "    ! [X1] :\n",
      "      ( X1 != butler\n",
      "     =&gt; hates(agatha,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_7),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_5,axiom,(\n",
      "    ! [X1,X2] :\n",
      "      ( killed(X1,X2)\n",
      "     =&gt; ~ richer(X1,X2) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_5),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_8,axiom,(\n",
      "    ! [X1] :\n",
      "      ( ~ richer(X1,agatha)\n",
      "     =&gt; hates(butler,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_8),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_6,axiom,(\n",
      "    ! [X1] :\n",
      "      ( hates(agatha,X1)\n",
      "     =&gt; ~ hates(charles,X1) ) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_6),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_10,axiom,(\n",
      "    ! [X1] :\n",
      "    ? [X2] : ~ hates(X1,X2) ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_10),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "fof(pel55_11,axiom,(\n",
      "    agatha != butler ),\n",
      "    file('/export/starexec/sandbox/benchmark/theBenchmark.p',pel55_11),\n",
      "    [interesting(-1.00)]).\n",
      "\n",
      "% AGInTRater output end\n",
      "\n",
      "% END OF SYSTEM OUTPUT\n",
      "% RESULT: SOT_nr4rgR - AGInTRater---0.0 says Rated - CPU = 0.00 WC = 0.09 \n",
      "% OUTPUT: SOT_nr4rgR - AGInTRater---0.0 says Derivation - CPU = 0.00 WC = 0.09 \n",
      "\n"
     ]
    }
   ],
   "source": [
    "print(text[begin:end])"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": []
  }
 ],
 "metadata": {
  "kernelspec": {
   "display_name": "Python 3",
   "language": "python",
   "name": "python3"
  },
  "language_info": {
   "codemirror_mode": {
    "name": "ipython",
    "version": 3
   },
   "file_extension": ".py",
   "mimetype": "text/x-python",
   "name": "python",
   "nbconvert_exporter": "python",
   "pygments_lexer": "ipython3",
   "version": "3.7.0"
  }
 },
 "nbformat": 4,
 "nbformat_minor": 2
}
